```python
# Initialize Otter
import otter
grader = otter.Notebook("Final_Exam.ipynb")
```

# ⌛️ Final Exam

This Final Exam is designed to test your mastery of the Python programming language. You will be asked to respond to a series of questions both multiple choice and short answer problems using python.


## Instructions

* You are allowed to use anything in jupyterlab to help you solve the problems. 
* You cannot open any other jupyter notebooks during the exam.
* You are not allowed to open any additional files. 
* You must run all code and press the submit button when shown to record your grades. If you do not do this your scores might not be recorded. 
* All of your responses should be persistent if you restart the kernel, however entries can only be updated if you run the cell and hit submit. 

**The exam is out of 238 points**


## Part 1: Entering Your Information for Credit

To receive credit for assignments, it is important we can identify your work from others. To do this, we will ask you to enter your information in the following code block.

You must hit submit and see the message saying the data is saved. If you do not see this message contact your instructor to ensure you are getting credit for your work.

**Note: do not delete the output.log file or the .responses.json file or you might lose your current work.**

### Before you begin

1. Run the block of code at the top of the notebook that imports and sets up the autograder. This will allow you to check your work. 
2. Run the block of code below to ensure your packages are up to date

All of the grader tests check to ensure your information is saved and accessible by the grader. If you see an error stating that your information is not found please restart the kernel and enter your student information again. 

```python
import pkg_resources
from subprocess import call
import sys

package_name = 'ENGR131_Util_2024'
version = '1.0.13'
package_version = f'{package_name}=={version}'

try:
    # Check if the package and version are installed
    pkg_resources.require(package_version)
    print(f'{package_version} is already installed.')
except pkg_resources.DistributionNotFound:
    # If not installed, install the package
    print(f'{package_version} not found. Installing...')
    call([sys.executable, '-m', 'pip', 'install', package_version])
except pkg_resources.VersionConflict:
    # If a different version is installed, you can choose to upgrade/downgrade
    installed_packages = {dist.key: dist.version for dist in pkg_resources.working_set}
    installed_version = installed_packages.get(package_name.lower())
    print(f'{package_name} {installed_version} is installed, but {version} is required.')
    # Optionally, upgrade or downgrade the package to the required version
    call([sys.executable, '-m', 'pip', 'install', '--upgrade', package_version])
```

```python
from ENGR131_Util_2024 import responses, StudentInfoForm, cell_logger
from ENGR131_Util_2024 import submit_question, ResponseStore, ValidateLogFile


# Register the log function to be called before any cell is executed
get_ipython().events.register('pre_run_cell', cell_logger)
responses["assignment"] = "final_1"

StudentInfoForm(**responses)
```

**Question 1: Match the Correct Definition**

*18 points*

Please run the block of code below, and respond to the prompts with the most suitable response. 

Make sure you hit the submit button or your grade will not be recorded. Confirmation that your responses are recorded will appear below the code block. We will only grade your most recent set of responses. You can submit entries as many times as you like. 

```python
from ENGR131_Util_2024 import responses
from ENGR131_Util_2024 import TypesQuestion_f2024 as TypesQuestion

TypesQuestion(**responses)
```

**Question 2: Select the option that matches the definition** 

*8 points*

Please run the block of code below, and respond to the prompts with the most suitable response. 

Make sure you hit the submit button or your grade will not be recorded. Confirmation that your responses are recorded will appear below the code block. We will only grade your most recent set of responses. You can submit entries as many times as you like. 

```python
from ENGR131_Util_2024 import MCQQuestion_f2024 as MCQQuestion
from ENGR131_Util_2024 import responses

MCQQuestion(**responses)
```

**Question 3: Select all statements which are TRUE** 

*20 points*

Please run the block of code below, and respond to the prompts with the most suitable response. 

Make sure you hit the submit button or your grade will not be recorded. Confirmation that your responses are recorded will appear below the code block. We will only grade your most recent set of responses. You can submit entries as many times as you like. 

**Note** If text is not fully rendered (e.g, you see at the end of the sentance ...) zoom out with control/cmd + -  

```python
from ENGR131_Util_2024 import responses
from ENGR131_Util_2024 import SelectMany_f2024 as SelectMany


out = SelectMany(**responses)
```

**Question 4: Reading and Interpreting Python Code** 

*45 points*

You are provided with the following code:

```python
1 numbers = [5, 4.0]
2 total = 0
3 
4 for num in numbers:
5    while total < 9:
6       total += num + 1.0
7    break
```

Please run the block of code below and respond to the prompts with the most suitable response. 

Make sure you hit the submit button or your grade will not be recorded. Confirmation that your responses are recorded will appear below the code block. We will only grade your most recent set of responses. You can submit entries as many times as you like. 

1. You will be asked to select the most relevant comment for the lines of code selected. 
2. You will be asked to describe how the code is executed. Please use the debugger tool to help ensure you have a correct answer.

   1. Line number: This is the line number of the code that is run during the execution of the code.
   2. Variable Changed: This is the variable whose value is changed during the execution of the code. 

      **Note** if no variable is changed - such as when evaluating logical expressions, select "None".
   3. Current Value: This is the value of the variable after the line of code is executed. 
      * If a logical expression is evaluated, please enter the value of the logical expression. 
      * If a control statement is run enter N/A.
   4. DataType: This is the data type of the variable after the line of code is executed. 
      * If a logical expression is evaluated, please enter the data type of the logical expression. 
      * If a control statement is run enter N/A. 
   
      **Note**: The autograder has mechanisms to provide partial credit. If you skip a step, it will continue grading at the next correct step. 

We provide the code in Python for your convenience. Do not alter the code to avoid responding incorrectly.

```python
numbers = [5, 4.0]
total = 0

for num in numbers:
   while total < 9:
      total += num + 1.0
   break
```

```python
from ENGR131_Util_2024 import ReadingPythonQuestion_f2024 as ReadingPythonQuestion
from ENGR131_Util_2024 import responses

ReadingPythonQuestion(**responses)
```

## Recommended

We recommend that you run the following code to ensure your responses are recorded. This is an extra measure to ensure your results saved. This can take a minute to run.

You can interrupt this execution by pressing the stop button. 

```python
submit_question()
```

## Question 5: Using Data Types 

*30 Points*

We have written a function that converts a DataType. Pass values with the expected DataTypes as defined below to the `convert_to_string` function. The function will print the output.

0. Run the `convert_to_string` function first.
1. Pass an integer `57` to `convert_to_string`.
    - Save the output to a variable `int_`.
2. Pass a float `6.28318` to `convert_to_string`.
    - Save the output to a variable `float_`.
3. Pass a boolean with a false value to `convert_to_string`.
    - Save the output to a variable `bool_`.
4. Pass a string `Python is fun` to `convert_to_string`.
    - Save the output to a variable `string_`.
5. Pass a list of integers `4`, `5`, `6` to `convert_to_string`.
    - Save the output to a variable `list_`.
6. Pass a tuple of strings `"d"`, `"e"`, `"f"` to `convert_to_string`.
    - Save the output to a variable `tuple_`.
7. Pass a dictionary with a key `color` and a value `blue` as a string and a key `number` and a value `8` as an integer to `convert_to_string`.
    - Save the output to a variable `dict_`.


```python
def convert_to_string(data):
    if isinstance(data, int):
        print(data)
    elif isinstance(data, float):
        print("{:.2f}".format(data))
    elif isinstance(data, bool):
        print(str(data))
    elif isinstance(data, str):
        print(data)
    elif isinstance(data, list):
        elements = ", ".join([str(x) for x in data])
        print(elements)
    elif isinstance(data, tuple):
        elements = ", ".join([str(x) for x in data])
        print(elements)
    elif isinstance(data, dict):
        pairs = ", ".join(
            ["Key = " + str(k) + " Value = " + str(v) for k, v in data.items()]
        )
        print(pairs)
    return data
```

```python
# 1. Pass an integer `57` to `convert_to_string`. 
# Save the output to a variable `int_`.
# completion of this step will pass test 1
...
```

```python
# 2. Pass an float `6.28318` to `convert_to_string`. 
# Save the output to a variable `float_`.
# completion of this step will pass test 2
...
```

```python
# 3. Pass a boolean with a false value to `convert_to_string`. 
# Save the output to a variable `bool_`.
# completion of this step will pass test 3
...
```

```python
# 4. Pass an string `Python is fun` to `convert_to_string`. 
# Save the output to a variable `string_`.
# completion of this step will pass test 4
...
```

```python
# 5. Pass a list of integers `4`, `5`, `6` to `convert_to_string`. 
# Save the output to a variable `list_`.
# completion of this step will pass test 5
...
```

```python
# 6. Pass a tuple of strings "d", "e", "f" to `convert_to_string`. 
# Save the output to a variable `tuple_`.
# completion of this step will pass test 6
...
```

```python
# 7. Pass a dictionary with a key `color` and a value `blue` as a string and a key `number` and a value `8` as an integer to `convert_to_string`. 
# Save the output to a variable `dict_`.
# completion of this step will pass test 7
...
```

```python
grader.check("q5 - Exploring DataTypes")
```

## Recommended

We recommend that you run the following code to ensure your responses are recorded. This is an extra measure to ensure your results saved. This can take a minute to run.

You can interrupt this execution by pressing the stop button. 

```python
submit_question()
```

## Question 6: Building a Stress Sensor

*36 Points*

In mechanical engineering, it is common to monitor the stress experienced by materials under load. We will build a class to implement a stress sensor. This sensor will save and also set a stress value in Pascals. Additionally, the class will convert the stress value given in Pascals to pounds per square inch (PSI).

1. Build a class called `StressSensor`. This class will have an initialization method and two additional methods.
    - Add an initialization method that accepts an initial stress value in Pascals as a variable `stress`.
      - Save the `stress` as an object attribute `stress_pa`. This is the stress in Pa.

    - Build a method `set_stress` that will set the object attribute `stress_pa`.
      - This method takes one input `stress`.
      - We want to make sure the input is a number. Do this by **trying** to convert it to a `float` using the built-in `float` method.
        - If `stress` is a number that can be converted to a float, save it to the object attribute `stress_pa`.
       - If `stress` cannot be converted to a float, set the object attribute `stress_pa` to the string 'fail'.

    - Build a method `to_psi` that converts the stress from Pascals to pounds per square inch.
      - Convert the stress using the following equation.
        $$PSI = Pascals \times 0.000145038$$
      - Return the stress in PSI.  

2. Instantiate a `StressSensor` `material_sensor` with an initial `stress` of `85000` Pa.
   
3. Extract the stress in Pa from the object `material_sensor`.
    - Assign the output to a variable `stress_pa`.

4. Extract the stress in PSI from the object `material_sensor`.
    - Assign the output to a variable `stress_psi`.

```python
# Your class for StressSensor goes here
...

# Instantiate the StressSensor material_sensor here
...

# extract the stress in Pascals
...

# extract the stress in pounds per square inch
...
```

```python
grader.check("q6-Building a Stress Sensor")
```

## Recommended

We recommend that you run the following code to ensure your responses are recorded. This is an extra measure to ensure your results saved. This can take a minute to run.

You can interrupt this execution by pressing the stop button. 

```python
submit_question()
```

## Question 7: Determining a Price for Structural Beams 

*40 Points*

As a structural engineer, you might want to estimate the total price at which you would sell a set of structural beams you've designed based on the cost of materials from a supplier and a desired profit margin. We will write a Python function to assist you in this task.

We have provided you with two dictionaries, `material_cost` and `beam_requirements`. `material_cost` defines the costs of materials from the supplier, and `beam_requirements` defines the materials needed to construct the structural beams. The structure includes two `Custom Alloy Beams`, which are not available from the supplier, that each costs $1200.

*Cost of Materials from Supplier*

| Material Name (Key) | Cost per Unit (value, in $) |
| -------------------- | --------------------------- |
| Steel Beam  | 250                           |
| Aluminum Beam | 150                          |
| Reinforcement Bar   | 50                         |
| Concrete Mix  | 120                          |

*Materials Required for Building the Structural Beams*

| Material Name (Key)       | Quantity (sub-key) | Cost to Build (value, in $)|
| ---------------------- | --------------- | --------------------------- |
| Steel Beam             | 10              | N/A                         |
| Concrete Mix           | 3               | N/A                         |
| Reinforcement Bar      | 7               | N/A                         |
| Custom Alloy Beam      | 2               | 1200                        |

1. Import the package `numpy` as `np`.
2. Define a function `price_calculator` that accepts three inputs: `material_cost`, `beam_requirements`, and `profit`.
    - Initialize a variable `cost_of_materials`, and set it equal to 0.
    - Using a loop, compute the total `cost_of_materials`.
        - For standard materials, fetch the cost per unit from `material_cost`.
        - For custom components like `Custom Alloy Beam`, use the cost specified in `beam_requirements`.
    - Determine the `sales_price` by applying the profit margin to the `cost_of_materials`.
        - $$sales\_price = cost\_of\_materials \times (1 + profit)$$
        - Round the final sales price to 2 decimal places using `np.round`.
    - Call the provided `print_total` function to get a string to print.
    - Print and return the final `sales_price`. 
3. Call the function `price_calculator` with the appropriate inputs and save the output to a variable `total_price`.

```python
# Your import statement goes here.
...

material_cost = {
    "Steel Beam": 250,
    "Aluminum Beam": 150,
    "Reinforcement Bar": 50,
    "Concrete Mix": 120,
}

beam_requirements = {
    "Steel Beam": {"quantity": 10},
    "Concrete Mix": {"quantity": 3},
    "Reinforcement Bar": {"quantity": 7},
    "Custom Alloy Beam": {"quantity": 2, "price": 1200},
}

# profit assignment goes here, we have set this to 63%
profit = 0.63

# We have provided you with a function to print the total cost of the order
# you need to call this function but do not need to change it.
def print_total(sales_price):
    return f"Total cost of the order: ${sales_price}"

# Your function price_calculator goes here
# Completion of this step should pass test 3
...

# Your call of the function price_calculator goes here
# Completion of this step should pass all tests
...
```

```python
grader.check("q7-Determining a Price of a Structural Beam")
```

## Recommended

We recommend that you run the following code to ensure your responses are recorded. This is an extra measure to ensure your results saved. This can take a minute to run.

You can interrupt this execution by pressing the stop button. 

```python
submit_question()
```

## Question 8: Plotting and Fitting (41 Points):

A prevalent task in engineering involves fitting data to a specific model. This exercise will focus on generating noisy data influenced by a sinusoidal function, a fundamental pattern observed in many engineering and physical phenomena. We will plot this data and perform a fitting procedure to align it with our sinusoidal model.

1. Begin by importing the necessary libraries. If you've already imported these in a previous part of your work, you might not need to repeat this step; however, it's generally a good practice to ensure all required packages are available.

   - Import `numpy` and alias it as `np`.
   - Import the `pyplot` module from `matplotlib`, aliasing it as `plt`.
   - Import the `curve_fit` function from `scipy.optimize`.

2. Define a function `sinusoidal` to calculate and return the values of a sinusoidal function based on the equation:
   $$A \sin(Bx + C)$$
   - This function should accept inputs `x`, `A`, `B`, and `C`.

3. Utilize the `linspace` function in `numpy` to generate a sequence from $-\pi$ to $\pi$ comprising `100` steps. 
   - Assign this sequence to the variable `x_data`.

4. Apply the `sinusoidal` function over `x_data` to calculate the true values when `A` is `2.5`, `B` is `1.2`, and `C` is `0.5`. 
   - Store these true values in `y_true`.

5. Generate noisy data by adding Gaussian-distributed noise to `y_true` using the `random.normal` method in `numpy`. 
   - Ensure the noise is of the same dimension as `x_data`. 
   - Scale the noise by `0.2` before adding it to `y_true`, and save the resultant noisy data in `y_noisy`.
   - **Hint** you can view the docstring by running `??np.random.normal`

6. Plot the noisy data (`x_data` vs `y_noisy`) using the `scatter` function from `matplotlib.pyplot`. 
   - Label the plot as "Noisy Data". 
   - Assign this plot to `plot_1`.

7. Employ the `curve_fit` function to align the noisy data with the `sinusoidal` function across the `x_data` range. 
   - Store the optimal fitting parameters in `popt` and the covariance of the parameters in `pcov`. 
     - **Avoid providing any initial guess** or use `None` to ensure compatibility with automatic grading systems.

8. Compute the fitted curve using the `sinusoidal` function with the optimal parameters found in the previous step. 
   - The unpacking operator (`*`) can be used to apply the list of fitted parameters directly. 
   - Save this fitted curve data in `y_fit`.

9. Plot the fitting results (`x_data` vs `y_fit`) using the `plot` function in `pyplot`, setting the line color to green with the `g` marker.
    - Label this plot as "Fitted Sinusoidal". 
    - Store this plot as `plot_2`.

10. Set the `xlabel` and `ylabel` of the plot to `x` and `y`, respectively, using the appropriate `plt` methods.

11. Finally, add a legend to the plot with the `legend` method in `plt` to distinguish between the noisy data and the fitted curve.

```python
# 1. Add your import statements here. 
# Completion of this step will pass the first test
...

# do not delete the line of code below, it will cause problems with the autograder.
np.random.seed(42)

# 2. Define the Sinusoidal function.
# completion of this step will pass test 2.
...

# 3. Define your x_data vector
# completion of this step will pass test 3.
...

# 4. Call the sinusoidal function as described above
# save the output to the variable y_true
# completion of this step will pass test 4.
...

# 5. Generate the noisy data described above
# save the output to the variable y_noisy
# completion of this step will pass test 5.
# use the ?? module to view the docstring.
...

# 6. Plot the noisy data using a scatterplot
# make sure to add the label "Noisy Data"
# save the output to the variable plot_1
...

# 7. Fit the noisy data to the sinusoidal function using curve_fit
# save the outputs to the variables popt and pcov
# completion of this step will pass test 6.
...

# 8. Calculate the fitted curve using the sinusoidal function
# save the output to the variable y_fit
# completion of this step will pass test 7. 
...

# 9. Plot the fitted results as described in the instructions
# save the output to the variable plot_2
# make sure to add the label "Fitted Sinusoidal"
...

# 10. Add axis labels to the plot
# completion of this step will pass a test.
...

# 11. Add a legend to the plot
# completion of this step will pass all tests.
...
```

```python
grader.check("q8-Plotting and Fitting")
```

## Recommended

We recommend that you run the following code to ensure your responses are recorded. This is an extra measure to ensure your results saved. This can take a minute to run.

You can interrupt this execution by pressing the stop button. 

```python
submit_question()
```

## Validating your log file

Before your submit your results we would recommend running the following command to validate your log file. This will ensure that you have followed the correct format and that your log file is valid. 

```python
ValidateLogFile("./output.log", responses["assignment"], 4, question_max_scores={1: 18, 2: 8, 3: 20, 4:45, 5: 30, 6: 36, 7: 40, 8: 41})
```

## Submitting Your Assignment

To submit your assignment please use the following link the assignment on GitHub classroom.
   
Use this [link](https://classroom.github.com/a/-WUNDMbG) to navigate to the assignment on GitHub classroom.

**Please submit just your output.log file. No other files are needed.**


If you need further instructions on submitting your assignment, please ask your TA during lab. 

## Viewing your score

Each `log` file you have uploaded will have a file with the name of your file + `Grade_Report.md`. You can view this file by clicking on the file name. This will show you the results of the autograder. 

**Make sure you run your code, and run the tests. If you do not run the test you will not get credit for your work.**

```{note}
In python and particularly jupyter notebooks it is common that during testing you run cells in a different order, or run cells and modify them. This can cause there to be local variables needed for your solution that would not be recreated on running your code again from scratch. Your assignment will be graded based on running your code from scratch. This means before you submit your assignment you should restart the kernel and run all cells. You can do this by clicking `Kernel` and selecting `Restart and Run All`. If you code does not run as expected after restarting the kernel and running all cells it means you have an error in your code. 
```

## Fin

Tests for q5 - Exploring DataTypes (points: None)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import json
>>> responder = ResponseStore()
>>> points_ = [3, 3, 3, 3, 6, 6, 6]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q5_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q5_1'
>>> max_score = 3
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABl6L-pDv5p5hDbvPowuS3KmnJvavbodPIc6Rv2MYr6UaSty95p2RUWtzhd0Uta8PZVymdB7qFcP-1Ff4JgWgc_9FKP4w==')
>>> if str(int_) == out and isinstance(int_, int):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': int_}
>>> responder.add_response(response)
>>> assert isinstance(int_, int)
>>> assert str(int_) == out, 'int_ incorrectly implemented.'
```
failure_message: int_ incorrectly implemented.
success_message: int_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import json
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q5_2'
>>> max_score = 3
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABl6L_-fuzaC571cBwKwRoYf4-E3ZG8cXgWTrvThoyyMa59OYClsYkyqt77aXm9s_BtEpbV-Ag85aWIGic6Dcn8TAlo2A==')
>>> if str(float_) == out and isinstance(float_, float):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': float_}
>>> responder.add_response(response)
>>> assert isinstance(float_, float)
>>> assert str(float_) == out, 'float_ incorrectly implemented.'
```
failure_message: float_ incorrectly implemented.
success_message: float_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import json
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q5_3'
>>> max_score = 3
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABl6MCsq3JZUW-PsgT0me3n_J_YHZJTH90hSi8_wOMvR6V0uhdd0Z4ObqizkbE_3tl0jG1AfjgheMaEIH04ffH97vZ4Yw==')
>>> if str(bool_) == out and isinstance(bool_, bool):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': bool_}
>>> responder.add_response(response)
>>> assert isinstance(bool_, bool)
>>> assert str(bool_) == out, 'bool_ incorrectly implemented.'
```
failure_message: bool_ incorrectly implemented.
success_message: bool_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import json
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q5_4'
>>> max_score = 3
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABl6MDN40KGTUiCYSgq5JUH8h32I85QvQV1PBGYO_O4fXyVRdGaQc7nwr-iD7nqD0_W0pTUlkZahLzitwGr-WK6OHX8xA==')
>>> if str(string_) == out and isinstance(string_, str):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': string_}
>>> responder.add_response(response)
>>> assert isinstance(string_, str)
>>> assert str(string_) == out, 'string_ incorrectly implemented.'
```
failure_message: string_ incorrectly implemented.
success_message: string_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import json
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q5_5'
>>> max_score = 6
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABl6MGrPcrPs3yo-J9Ih38ykNusMWQWZhWeSmhV71lIFjMtchzRh9kdnf4KHKyfX-JA3vaWtKIwLC01YUKdl8LaWgOhZg==')
>>> if str(list_) == out and isinstance(list_, list):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': list_}
>>> responder.add_response(response)
>>> assert isinstance(list_, list)
>>> assert str(list_) == out, 'list_ incorrectly implemented.'
```
failure_message: list_ incorrectly implemented.
success_message: list_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import json
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q5_6'
>>> max_score = 6
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABl6MI7e1k8IEJ87zyTLI26Rj6eH9wPwbRLzfBfcX9uNTxcg8qDNbTg9MaLp9dQJQrnYbNQpQINpgebXFFBqMRqXFviuQ==')
>>> if str(tuple_) == out and isinstance(tuple_, tuple):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': tuple_}
>>> responder.add_response(response)
>>> assert isinstance(tuple_, tuple)
>>> assert str(tuple_) == out, 'tuple_ incorrectly implemented.'
```
failure_message: tuple_ incorrectly implemented.
success_message: tuple_ correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import json
>>> responser = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q5_7'
>>> max_score = 6
>>> score = 0
>>> out = get_string_from_encrypted_string(b'gAAAAABl6MLBvYi3_vYH5B5AXz4hCVykYrvuRMCd9rzFwS2TnRpgWUXYF90NG_rSBOB4nKHNGh1t8LJeGow9L25DM0AxnNLkpwRgjQyW81CBlhCEYQtIw6Y=')
>>> if str(dict_) == out and isinstance(dict_, dict):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': dict_}
>>> responser.add_response(response)
>>> assert isinstance(dict_, dict)
>>> assert str(dict_) == out, 'dict_ incorrectly implemented.'
```
failure_message: dict_ incorrectly implemented.
success_message: dict_ correctly implemented.


Tests for q6-Building a Stress Sensor (points: None)
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import json
>>> responder = ResponseStore()
>>> points_ = [6, 6, 8, 8, 8]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q6_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q6_1'
>>> max_score = 6
>>> score = 0
>>> if isinstance(StressSensor, type):
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': isinstance(StressSensor, type)}
>>> responder.add_response(response)
>>> assert isinstance(StressSensor, type), 'The Class StressSensor is not implemented.'
```
failure_message: The Class StressSensor is not implemented.
success_message: The Class StressSensor is implemented.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q6_2'
>>> max_score = 6
>>> score = 0
>>> t = StressSensor(25000)
>>> if hasattr(t, 'stress_pa') and t.stress_pa == 25000:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': t.stress_pa}
>>> responder.add_response(response)
>>> assert hasattr(t, 'stress_pa'), "StressSensor does not have 'stress_pa' attribute"
>>> assert t.stress_pa == 25000, 'StressSensor initialization did not set attribute stress_pa correctly'
```
failure_message: The Class StressSensor initialization is implemented incorrectly.
success_message: The Class StressSensor initialization is implemented correctly.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q6_3'
>>> max_score = 8
>>> score = 0
>>> t = StressSensor(25000)
>>> t.set_stress(50000)
>>> if t.stress_pa == 50000.0:
...     score += 4
>>> t_ = StressSensor(25000)
>>> t_.set_stress('abc')
>>> out = str(t_.stress_pa)
>>> if out == 'fail':
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': [t.stress_pa, t_.stress_pa]}
>>> responder.add_response(response)
>>> t = StressSensor(25000)
>>> t.set_stress(50000)
>>> assert t.stress_pa == 50000.0, 'set_stress does not function properly'
>>> t = StressSensor(25000)
>>> t.set_stress('abc')
>>> out = str(t.stress_pa)
>>> assert out == 'fail', 'set_stress does not return the correct response when a unsuitable input is provided'
```
failure_message: set_stress can not correctly handle non numeric inputs.
success_message: set_stress can correctly handle non numeric inputs.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import json
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q6_4'
>>> max_score = 8
>>> score = 0
>>> t = StressSensor(68947.57)
>>> if abs(t.to_psi() - 10.0) < 0.01:
...     score += 4
>>> hold = t.to_psi()
>>> t.set_stress(137895.14)
>>> if abs(t.to_psi() - 20.0) < 0.01:
...     score += 4
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': [hold, t.to_psi()]}
>>> responder.add_response(response)
>>> t = StressSensor(68947.57)
>>> assert abs(t.to_psi() - 10.0) < 0.01, 'Incorrect conversion to psi.'
>>> t.set_stress(137895.14)
>>> assert abs(t.to_psi() - 20.0) < 0.01, 'Incorrect conversion to psi.'
```
failure_message: Incorrect conversion to psi.
success_message: Correct conversion to psi.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> import json
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q6_5'
>>> max_score = 8
>>> score = 0
>>> stress_sensor = StressSensor(68947.57)
>>> stress_pa = stress_sensor.stress_pa
>>> stress_psi = stress_sensor.to_psi()
>>> conditions = [hasattr(stress_sensor, 'stress_pa'), hasattr(stress_sensor, 'to_psi'), stress_pa == 68947.57, abs(stress_psi - 10.0) < 0.01]
>>> statements = ["StressSensor does not have 'stress_pa' attribute", "StressSensor does not have 'to_psi' method", 'Incorrect stress value for stress_pa.', 'Incorrect conversion value for stress_psi.']
>>> for condition in conditions:
...     if condition:
...         score += 2
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': conditions}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: Incorrect conversion to psi.
success_message: Correct conversion to psi.


Tests for q7-Determining a Price of a Structural Beam (points: None)
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> import numpy as np
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> points_ = [0, 1, 0, 6, 15, 12, 6]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q7_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q7_1'
>>> max_score = 0
>>> score = 0
>>> conditions = [material_cost['Steel Beam'] == 250, material_cost['Aluminum Beam'] == 150, material_cost['Reinforcement Bar'] == 50, material_cost['Concrete Mix'] == 120]
>>> statements = ['Steel Beam not correctly defined', 'Aluminum Beam not correctly defined', 'Reinforcement Bar not correctly defined', 'Concrete Mix not correctly defined']
>>> for condition in conditions:
...     if condition:
...         score += 0
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': conditions}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: material_cost is incorrect. Please make sure you did not modify the dictionary.
success_message: material_cost is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> import json
>>> responser = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q7_2'
>>> max_score = 1
>>> score = 0
>>> if np.__version__ is not None:
...     score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': np.__version__}
>>> responser.add_response(response)
>>> assert np.__version__ is not None, 'numpy is incorrectly imported'
```
failure_message: numpy incorrectly implemented.
success_message: numpy correctly implemented.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q7_3'
>>> max_score = 0
>>> score = 0
>>> conditions = [beam_requirements['Steel Beam']['quantity'] == 10, beam_requirements['Concrete Mix']['quantity'] == 3, beam_requirements['Reinforcement Bar']['quantity'] == 7, beam_requirements['Custom Alloy Beam']['quantity'] == 2, beam_requirements['Custom Alloy Beam']['price'] == 1200]
>>> statements = ['Steel Beam quantity not correctly defined', 'Concrete Mix quantity not correctly defined', 'Reinforcement Bar quantity not correctly defined', 'Custom Alloy Beam quantity not correctly defined', 'Custom Alloy Beam price not correctly defined']
>>> for condition in conditions:
...     if condition:
...         score += 0
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': conditions}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: beam_requirements is incorrect. Please make sure you did not modify the dictionary.
success_message: beam_requirements is correct.
```
>>> import drexel_jupyter_logger
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import responses
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q7_4'
>>> max_score = 6
>>> score = 0
>>> conditions = [price_calculator.__name__ == 'price_calculator', price_calculator.__code__.co_argcount == 3]
>>> statements = ['Function name is incorrect. It should be price_calculator.', 'The price_calculator function does not have the correct number of arguments (expected: material_cost, beam_requirements, profit).']
>>> for condition in conditions:
...     if condition:
...         score += 3
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': [price_calculator.__name__, price_calculator.__code__.co_argcount]}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: price_calculator call is incorrect. Ensure function name, arguments, and calculation logic are correct for structural beams.
success_message: price_calculator call is correct for structural beam price estimation.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import responses
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q7_5'
>>> max_score = 15
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABl6MvajDvDxG6mf-GXt8a7PK8yxNRY5RvjDyB-Uhi2Pe1QxOTDks-IcS0NlqA8QRtVEvN-mDzPusk5IyJa5Mr70LqFeQ==')
...     if str(price_calculator(material_cost, beam_requirements, 0.2)) == out:
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': out}
>>> responder.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     assert str(price_calculator(material_cost, beam_requirements, 0.2)) == out, 'price_calculator  computation is incorrect.'
```
failure_message: price_calculator computation is incorrect.
success_message: price_calculator computation is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import responses
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q7_6'
>>> max_score = 12
>>> score = 0
>>> parts_2 = material_cost.copy()
>>> parts_2['Steel Beam'] = 10
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABl6MxU5vtTWbI4SgDP6bof8xgVvoiD7kLugrfSyWriqXlSa729QyUzs_Wa-cDOzS28ueLXLnV7ACX7aj_2MLu7kILdXw==')
...     if out == str(price_calculator(parts_2, beam_requirements, 0.2)):
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': out}
>>> responder.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     assert out == str(price_calculator(parts_2, beam_requirements, 0.2)), 'Your code does not use the beam_requirements correctly'
```
failure_message: Your code does not use the material_cost correctly.
success_message: Your code correctly uses the material_cost.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import responses
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q7_7'
>>> max_score = 6
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABl6Mznzrh3BqQgYjsfq_nggPhmz1ZkwL6rN_bNPRbdjazndqzFxOzzEV0AkkXvCh_PCxxc5rH2cMCo_wpbE7rXOnAwz7xiFWfpenZTeG4kchmdnQoNH0RJ9m2CLx1IpCqfGoJl')
...     if print_total(price_calculator(material_cost, beam_requirements, profit)) == out:
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': out}
>>> responder.add_response(response)
>>> with patch('builtins.print') as mock_print:
...     assert print_total(price_calculator(material_cost, beam_requirements, profit)) == out, 'The string being printed is not correct'
```
failure_message: print message is incorrect.
success_message: print message is correct.


Tests for q8-Plotting and Fitting (points: None)
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import responses
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> points_ = [2, 6, 4, 2, 3, 6, 2, 2, 2, 2, 2, 4, 4]
>>> for (i, point) in enumerate(points_):
...     drexel_jupyter_logger.variable_logger_csv(f'0, {point}', f'q8_{i + 1}')
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_1'
>>> max_score = 2
>>> score = 0
>>> conditions = [np.__name__ == 'numpy', plt.__name__ == 'matplotlib.pyplot', 'curve_fit' in dir()]
>>> statements = ['numpy not imported correctly', 'matplotlib.pyplot not imported correctly', 'scipy.optimize.curve_fit not imported correctly']
>>> for condition in conditions:
...     if condition:
...         score += 2 / len(conditions)
>>> score = int(round(score, 0))
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': plt.__name__}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: Your input statements are incorrect.
success_message: Your input statements are correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import responses
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_2'
>>> max_score = 6
>>> score = 0
>>> x_test = np.array([0, 1, 2])
>>> (a_test, b_test, c_test) = (1, 1, 1)
>>> expected_output = np.array([0.84147098, 0.90929743, 0.14112001])
>>> conditions = [np.allclose(sinusoidal(x_test, a_test, b_test, c_test), expected_output)]
>>> statements = ['sinusoidal function is not implemented correctly']
>>> for condition in conditions:
...     if condition:
...         score += 6
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': [*sinusoidal(x_test, a_test, b_test, c_test)]}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: sinusoidal function is not implemented correctly
success_message: sinusoidal function is implemented correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import responses
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_3'
>>> max_score = 4
>>> score = 0
>>> conditions = [np.allclose(x_data[0], -np.pi), np.allclose(x_data[-1], np.pi), len(x_data) == 100, np.isclose(sum(x_data[3::3]), 3.141592653589801)]
>>> statements = ['x_data is not defined correctly', 'x_data is not defined correctly', 'x_data length is not correct', 'x_data is not defined correctly']
>>> for condition in conditions:
...     if condition:
...         score += 1
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(conditions)}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: x_data is not defined correctly.
success_message: x_data is defined correctly.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import responses
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_4'
>>> max_score = 2
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABl6OsbZQjh0i9RMNKzHdS3XHchdX1y11dJFaOJUR2JCtIEHeL1Z1T1EzXG8R6Rh2FFBE6DueCEeQ-E9F5ZJHHQoXTSVZyfQGTjuHFUS4z9-i7v5RKOV08uYnuw_ILTI4DuWg-CvTfjp-DlQ22FPcSC9EBpcL_j8yAVh0qmS-yoKsCeldw9-XjW-YhvLDDImhz5bQRMXEMHzcL5cIkGcj3fKz2UwU2A9-uaj6ywdF1bFPCwF69PH8B9laY7h1f9DnBSbecDKgS1W7z8RwOmzgXfGoEDaQ8sVx8LjG2tgRVGY1D2dqeta1pU2TG1yeocOUflXlOCjnPidrnskGoiOQgpb08B6AFi81yjpzOCzokz261V_-wycXTcpEA8SXeDf9NxXtTpo66Bty6qZR49uuNchQFMts6AS8OypsXEI00CwDUAduRdr7VTlLWSbAjHICNu1tsM09jUvcz9-VjIlfv5lMD2U0uMqh-8X9n-jjg6ia8ZJbXSl_tSCGRYFjd50zsCTLBGVBVSo4BwQPL3UE5xLSUZdGKBxXtadw89Gz8IU5HXklkmbagG2phVABCLi-5nUrd30u-Q3CZXzV3ftig2SvTjzD6Rq9maIJAB5tpUDPaEUZRx08z4XK0b0KbbTseRqnN4atU8rQpwLSrcWtdxCYJ2e3TBSAmYV075VV3OAef7R-j1dN_ZYdKlcL3VA7GMtMjfZqfigGycG1Yj8vWpTgLuRgplHhgnkPSrciCwt9-EeDzGl76zBYd-j3veme-CaHhTxLWDUwauDHtMBy745xd9BQYiS6PcofHxVc0nygk39M5nMGIBeZqaNAqmovkQlIBAw92dFu-psEWPZxdJEpKuSurFJ-KWRvt3QUvHG4rvnM3IZAR4hqaC94-1RVL0L8LxtL7LY1Gun7RoyLDvRBNjVNcUqQT4YyyW0S9iiVoaJGCY-KMUp4A_fu-w2AcBuIg56l3KBjpxJRTROM_T122Vnaq7F_mLFycn3k6ruzt3uGJDRSXrvwAWPF_Mw1hxVDTobimaRuadpRbkLbtZAnxo1UqQ5c8zaCRBfTa6EWk-C7taKWRsqW-s_XB_DFgLMoJkasX7_8GstKCaEzY6SDXrrLNQ2V2WTHpPKzxjIZ8mrvty0F7a5S8rkvi-aWkFQmvuUesnyy7XKI2uAb-Jx97LBKkjWm9IHzcwh5B5B1r0mbaMfzB2PI15bBuesN-stCbbpvGgtSanvYTZ6ERik-uwK-NXUD1F2yK2dGQ0j2JlMTkXzW2dt4vZCEBLKd5-TmMkikLGU1PsLxIOjbgTDxW82WDl2CqIqARb2A8B-HneUT4sDvSmGylS7Z3_04rR5BnBX18kWWkdldsJWzjRTEMrVTCIxPSCLh1Y7v2_bUgZPJIqQlqw5hNwx4MYsmA2DigY45S5gYU1rXMVqvWUxr0lD-4RBtyL13KsOCMo8LcdbIWWTA4t0ur9UXnj2wY4FPoxub1BIEEi6gRcuQtbDdwG1ITcEdOSwkAn-dPWmc_kz88q71qLO5uydMAS9jCp_54n6M_2yudBeQQpGJT0lCklMtHD2XS0vADLxiyyd-2lC8QiPVECCTV2q8jSqTF0RPqg8Mx-MgEso0I2GDy4RXDJLTk3m3_IQYrLI9_jnV8lvLHfAaQ_YI8ZNUl4SPGvi8KPeWKgkefLwl2ZL1cp6Sw9hA==')
...     if str(y_true) == out:
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': str(y_true)}
>>> responder.add_response(response)
>>> assert str(y_true) == out, 'y_true is incorrect'
```
failure_message: y_true is incorrect.
success_message: y_true is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import responses
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_5'
>>> max_score = 3
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABl6OtEtOwzVaFP-LljEIHyP6-mr8OoF-fGBwX_Q3Sm4eFERLQQ-AyW0gLo2TNEB9fRZ3tpgddE8mlqGKFB-GLzHe0cEBiN_W-uSQzEpmfDL7uveLymqlJer8hz7ew0hZp1K5a2pjgSkEcroiLK5GE8aAmYGxpBTQ0aFV-6J5R1Dlr5siVApd-uYr_YEcDf2J7apeQAzO5cTz-shlM8u-aULYV8FFFf0AQt3kftySGKm2eEE2OlDJ3H9e6tvS3AaFbOpqos7DX1IJG3wTjiCEgI07oXMYBsObyerWtaxQCYp3TC_X3zfr4p1qiDxIHeTlrSgr8SinpS_pct8eidsO71J2HOAvRG_oJuDhUNqrm7bXzhA7mg0piB0LeeBmFH_jglzKXmX_sO3v_ixqvmEg5hUgrHxYulAcg_WhWNQp_fNiM9oXIjFdEKDZEvGjlOoQi3J2UTvKDc9UM1NDJRTftUeWOa6tatH9vmNMKUQgpT-xxiYB-a20038iSSMBna-rf8K5Pfz5YUwlCIuvsQoPPmgC-nosrs76lbSoau6XuHYPcBQNvnjK2OdKtGpDQbfcwnwRACIiWZIXjGDMLv7p9nxnEG4pZwFrQRoL470IkEM5qsCL_ZVd9lh3Qtn9Nr0qJLKbTD8nYJIkR0UpNNhMapbWJOpg9blozNkHASkvkR4bJjtuQRIwUE4pCknwjg8XkrgX1aw2tjQf_niSmVA25nAH5V-JCBHe2ZBRSDBiJb1Hw2Mhj8-2y38umnP_to_FYT3tIdTv5kstwxQ6UhgtWmYTMiZVee87WRQLuIjDSpDPPdwaUbwOmQyvdJ_WAbZCvMjI37U-RtR80yu86g7_Ss-YY0mj9z2GtsqKoxZmR51UztX91jHmPGJwrMLVN4655l1BtVG9DrOPiC_zldoFDEDhbC4WHfnb555hux7CDm_b3yh-PDmIdoJ6i9UBtTY-OJ9I7VpjqBg3mxAKFHDT2J52np30qoCAaHs_uBLv7xx3_BPFDEDgwsHlMkk0M3Nih6u0QeV3JQAFEjghpc_i1eJ4wTCHAUbTb3PIdmcROir7EwDCnewNOW9FtU5dTZnSESHzNgCt5i5b6UR7YVXuVqeEh3-I1BkUbSrXbVcHUoAkqnVYaqANS_hkkVPw9Cj85thgHM7ne2zGSozzREyztW6JUkf0Ew8LUKe3R6B5VJjpQxRIO87QiMHEG8AHKkbJFPhtcO3FeEClsVuLV7pMFVieC2ORe32dfqNtkSxx9cwtejS5FUsjfXWaaA_wJciF8eZBhBUHzoX0k-kKednL8afCva1UWVeom0PwCoBGe-lz5ES3quPKK1WW-kIDukrnxj2qkwTmbi47YrQ2BJaMB4xVw3Zb7PQMAmlk4QqDtDDdlsvwlii6e3i7ryYNBdmDfsDITdVCDl252STmbsMbDihdH9O_qzHm5mvyS8qq1FrFwnq30rl-kvpuXnjeX4hb1TKQTsX10tm3-3ltbTJIFl0tkS-7b3eLvdZJYCvu4Wb3aa_v_F9ifBVBaskjrSW4yr-bM1cnT6yK9NUA8xVZhS7TedLDY1xwhFNqqEQPWOrdZKIBUeDYXwlJziCpG8Zf3uYPsVRh92iGvwm_o8KabdHnE_BTMj6oWmAO-7HNdWr-zPAyCL5ZedCPn6rSkc7bPKcA_i6M_veZG9yF4UtSVnP_EQCw==')
...     if str(y_noisy) == out:
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': len(y_noisy)}
>>> responder.add_response(response)
>>> assert str(y_noisy) == out, 'y_noisy is incorrect'
```
failure_message: y_noisy is not implmemented correctly
success_message: y_noisy is implmemented correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import responses
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_6'
>>> max_score = 6
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABl6OuLOJTIDo4-g5gul9bHne9fZvoqMC-aRLfGuaW5jufR6c40uhOv9YNIdn3AJ4qaC8ep1fulysL5rCF57gj6wtHC76yB2ULQsk_xANFzO11m9iTI1Q5mBeO0oNmchIlV2_R_')
...     chars_to_remove = ['[', ']', '\n']
...     for char in chars_to_remove:
...         out = out.replace(char, '')
...     out = [float(x) for x in out.split(' ')]
...     if np.allclose(popt, out):
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': [*popt]}
>>> responder.add_response(response)
>>> assert np.allclose(popt, out), 'popt is not returning the correct values'
```
failure_message: popt is not returning the correct values
success_message: popt is returning the correct values
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import get_string_from_encrypted_string
>>> from unittest.mock import patch
>>> from ENGR131_Util_2024 import responses
>>> import json
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_7'
>>> max_score = 2
>>> score = 0
>>> with patch('builtins.print') as mock_print:
...     out = get_string_from_encrypted_string(b'gAAAAABl6OwH26bBtouW6OVlhUW0_kJ_k74V6vgjdhX7MgLzuxFAg8fJQGpd7zsHjOvmWM4ECLY1xAsyoQiGwShs8V9VHdZ9RjNLVEW3AkDSOye_fZOH4aOAEBvR_3R4vCYKRe0Q20jOO4KXmxQgJLg4_RiDM7j-m3LBAlERZRMlOF_Z9iYH2lZfX007jtizGFpmoorfn8Waa0xYTBCUmQ2OWT1tiTDfFdAZTkLERFeS4KFe7pcJ_mPAdSQJNUJ7jgbgR_zlEK-Bc8jeoVqsl1pZ23b1HD0VQykIs-dUydIjhZ3E7gTX0bNAHDNKU0t6U8qcYgtopBG9WnT4r8iGCJzmsVc-022XbV54eTuD2g5LTRC53GBMsNtarA91W-yu8uYWqNc7CcRJzbugLYcduDA7ez0KDuitaEbAA0H666ssYKTuoQ3YYBbLNDv0HDOY74X8EkQlBPPn1LNzXmVfbPCt4eOPInnREaDejmWnrHIqaCtGfcDS0YStN9PiTWU-HELMXDOZb0bFi1RGFZgjWXcIqZ6ggjkYU93Zxy03ZMvyHDoDNfc99oH9TBgefk7alU0EiQfrfwvzVRtJH9yuhP0Qfubuh2YFJw4j1vBtZflJvAP0azpb_pFUBMaoFXbmx8-XSITba1cDYoHl0V4knDiRXQiFNxe3-hQJX_JJJyCGIs7UPqaeSrbl77yYJQ33GcN5xHuTZbPibHTCzE10tTistmOYp3lVCHsxujiW5708h-yf7M-0u1ZPiixd5d0cfaZXmyRCrxGNb3Bvfm7Tfip-_PRzp8Coycj_vtm6uHX-xkfaoPzm6RqV-y1LBn_EdmeXL4AqKbuZ-HpcJbkgn7vd5XiGHZ6nAjnDBd58rT8uJIDDA50RLWw0xmkKHPOlmwB6a_7_iWV1MKW2620vBkikECIfamdWcnNUNT3aFVbl7TNsTF8r2k70YjyQTo4Pffl3CoVNC2eIXGBPgGTs0NLHOtq1XmKgrI-0lvBO__c2Qr72M1wa2mqqYjjcrHz5j5ZwzBKH3ayI4rvDfzrqAuBLyT0kkfwQt1H1wGbD82bt7ya3ZGYcJ-o62U4pFrsuKt2OyZUcTnG13tZycS9nGS_CRIhOJxy5gygU684TRecF5l0t1OWx0WNPWpv82EM_C6vAIasip2zVYmpuYgwDb_j7WWpFzSywjvJYU0F52Sow1QYJqkORIY5yKF88fxo6oudaLvykU96-SJE3jMMol_Wf9wQXxIGQkmEagjHuqXzN3gYgs1X1rkzbtJZJXH8wl_w7g-eHJZ_OwFldugolbJtGFBna2aFvPpteQl-v2rhClqGqZbLDOu3wIMEoF2ggC3VSx-bYKKh5J8jtc7oV-l5lFrUqHMzB3SRCe8pwAv6Dv7ptj6efOcUBKdTQ68t08WMTsKN-CEtUPhb6ZUnHGcI9Pr8mUTvAr6WyKzNMNIdFbqnwlQV1UolvunpesbItk6gtLeYcvfQ58bdGHjsn0Gbrq1U71P1F0qy_wiVX_brtXb2zXmcKBZalUgIU_BTHjEMkeTBcG01wllqOqBpwLlo2_AH10CDV0KcBU2HJU4BfaHot884O6m6nx9LY4WYkLuvXYDKpKB5WBi154v1U_BINtTjryQocGxZLpgy8m5kB3vFugCIcPk8OWUcnUqu_msPNBo1KFbxYx7M921j_aTE6uc796yoSUy9IxJ-4EveGigFD2E4hIAc=')
...     chars_to_remove = ['[', ']', '\n']
...     for char in chars_to_remove:
...         out = out.replace(char, '')
...     out = out.replace('  ', ' ')
...     out = [x for x in out.split(' ')]
...     out = [item for item in out if item != '']
...     out = [float(x) for x in out]
...     if np.allclose(y_fit, out):
...         score = max_score
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': [*popt]}
>>> responder.add_response(response)
>>> assert np.allclose(y_fit, out), 'y_fit is returning incorrect values'
```
failure_message: y_fit is returning incorrect values
success_message: y_fit is returning correct values
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import responses
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_8'
>>> max_score = 2
>>> score = 0
>>> conditions = [plot_1.axes.get_xlabel() == 'x', plot_1.axes.get_ylabel() == 'y']
>>> statements = ['xlabel is incorrect', 'ylabel is incorrect']
>>> for condition in conditions:
...     if condition:
...         score += 1
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': [plot_1.axes.get_xlabel(), plot_1.axes.get_ylabel()]}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: The plot labels are incorrect.
success_message: The plot labels are correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import responses
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_9'
>>> max_score = 2
>>> score = 0
>>> conditions = [plot_1.get_label() == 'Noisy Data', plot_2[0].get_label() == 'Fitted Sinusoidal']
>>> statements = ['Noisy Data', 'plot_1 is not labeled correctly', 'Fitted Sinusoidal', 'plot_2 is not labeled correctly']
>>> for condition in conditions:
...     if condition:
...         score += max_score / len(conditions)
>>> score = int(round(score, 0))
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': [plot_1.get_label(), plot_2[0].get_label()]}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: the plots are not labeled correctly
success_message: the plots are labeled correctly
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import responses
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_10'
>>> max_score = 2
>>> score = 0
>>> conditions = [plot_2[0].get_color() == (0.0, 0.5, 0.0, 1) or plot_2[0].get_color() == 'g' or plot_2[0].get_color() == 'green']
>>> statements = ['Fitted Sinusoidal fit is not the correct color']
>>> for condition in conditions:
...     if condition:
...         score += max_score / len(conditions)
>>> score = int(round(score, 0))
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': [plot_2[0].get_color()]}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: Fitted Sinusoidal fit is not the correct color
success_message: Fitted Sinusoidal fit is the correct color
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import responses
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_11'
>>> max_score = 2
>>> score = 0
>>> conditions = [plot_1.axes.get_legend() is not None]
>>> statements = ['The legend is not correct on the plot']
>>> for condition in conditions:
...     if condition:
...         score += 2
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': [plot_1.axes.get_legend() is not None]}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: The legend is not correct on the plot
success_message: The legend is correct on the plot
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import responses
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_12'
>>> max_score = 4
>>> score = 0
>>> conditions = [all(plot_2[0].get_xdata() == x_data), all(plot_2[0].get_ydata() == y_fit)]
>>> statements = ['the x data for plot_2 is not correct', 'the y data for plot_2 is not correct']
>>> for condition in conditions:
...     if condition:
...         score += max_score / len(conditions)
>>> score = int(round(score, 0))
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': conditions}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: plot_2 does not have the correct values.
success_message: plot_2 is correct.
```
>>> import drexel_jupyter_logger
>>> from ENGR131_Util_2024 import responses
>>> from ENGR131_Util_2024 import responses, ResponseStore
>>> responder = ResponseStore()
>>> if 'drexel_email' not in responses:
...     raise ValueError('Please fill out the student info form and run the test again')
>>> question_id = 'q8_13'
>>> max_score = 4
>>> score = 0
>>> conditions = [np.array_equal(plot_1.get_offsets()[:, 1], y_noisy), np.array_equal(plot_1.get_offsets()[:, 0], x_data)]
>>> statements = ['the y data for plot_1 is not correct', 'the x data for plot_1 is not correct']
>>> for condition in conditions:
...     if condition:
...         score += max_score / len(conditions)
>>> score = int(round(score, 0))
>>> drexel_jupyter_logger.variable_logger_csv(f'{score}, {max_score}', question_id)
>>> response = {'question_id': question_id, 'score': score, 'max_score': max_score, 'student_response': conditions}
>>> responder.add_response(response)
>>> for (condition, statement) in zip(conditions, statements):
...     assert condition, statement
```
failure_message: plot_1 does not have the correct values.
success_message: plot_1 is correct.
